A genome-wide association study implicates the olfactory system in Drosophila melanogaster diapause-associated lifespan extension and fecundity

Author response:

The following is the authors’ response to the previous reviews.

Reviewer #1 (Public Review):

Summary:

The paper begins with phenotyping the DGRP for post-diapause fecundity, which is used to map genes and variants associated with fecundity. There are overlaps with genes mapped in other studies and also functional enrichment of pathways including most surprisingly neuronal pathways. This somewhat explains the strong overlap with traits such as olfactory behaviors and circadian rhythm. The authors then go on to test genes by knocking them down effectively at 10 degrees. Two genes, Dip-gamma and sbb, are identified as significantly associated with post-diapause fecundity, and they also find the effects to be specific to neurons. They further show that the neurons in the antenna but not the arista are required for the effects of Dip-gamma and sbb. They show that removing the antenna has a diapause-specific lifespan-extending effect, which is quite interesting. Finally, ionotropic receptor neurons are shown to be required for the diapause-associated effects.

Strengths and Weaknesses:

Overall I find the experiments rigorously done and interpretations sound. I have no further suggestions except an ANOVA to estimate the heritability of the post-diapause fecundity trait, which is routinely done in the DGRP and offers a global parameter regarding how reliable phenotyping is.

We added to the Methods: “We performed a one-way ANOVA to get the mean squares for between-group and withingroup variances and calculated broad-sense heritability using the formula: H2 = MSG - MSE / MSG + (k-1) MSE where MSG - Mean square between groups and MSG - Mean square within groups and k - Number of individuals per group. Using this formula, the broad-sense heritability for normalized post-diapause fecundity was found to be 0.51.”

We added to the Results: “The broad-sense heritability for normalized post-diapause fecundity was found to be 0.51 (see Methods).”

A minor point is I cannot find how many DGRP lines are used.

Response: We screened 193 lines and have added that to the Results.

Reviewer #2 (Public Review):

Summary

In this study, Easwaran and Montell investigated the molecular, cellular, and genetic basis of adult reproductive diapause in Drosophila using the Drosophila Genetic Reference Panel (DGRP). Their GWAS revealed genes associated with variation in post-diapause fecundity across the DGRP and performed RNAi screens on these candidate genes. They also analyzed the functional implications of these genes, highlighting the role of genes involved in neural and germline development. In addition, in conjunction with other GWAS results, they noted the importance of the olfactory system within the nervous system, which was supported by genetic experiments. Overall, their solid research uncovered new aspects of adult diapause regulation and provided a useful reference for future studies in this field.

Strengths:

The authors used whole-genome sequenced DGRP to identify genes and regulatory mechanisms involved in adult diapause. The first Drosophila GWAS of diapause successfully uncovered many QTL underlying post-diapause fecundity variations across DGRP lines. Gene network analysis and comparative GWAS led them to reveal a key role for the olfactory system in diapause lifespan extension and post-diapause fecundity.

Comments on revised version:

While the authors have addressed many of the minor concerns raised by the reviewers, they have not fully resolved some of the key criticisms. Notably, two reviewers highlighted significant concerns regarding the phenotype and assay of post-diapause fecundity, which are critical to the study. The authors acknowledged that this assay could be confounded by the 'cold temperature endurance phenotype,' potentially altering the interpretation of their results.

However, they responded by stating that it is not obvious how to separate these effects experimentally. This leaves the analysis in this research ambiguous, as also noted by Reviewer #3.

We should have clarified earlier that we actually chose to measure post-diapause fecundity in order to minimize any impact of ‘cold temperature endurance.” In fact, we chose post-diapause fecundity as the appropriate measure of successful diapause for both technical and conceptual reasons. Conceptually, the benefit of diapause is to perpetuate the species. It seems obvious to us that post-diapause fecundity is more relevant to species propagation than other measures of diapause such as how many egg chambers contain yolk or how many eggs are laid. Technically, we chose 5-week diapause and recovery based on pilot studies that showed that nearly all DGRP lines showed excellent survival at 5 weeks in diapause conditions. Therefore, our experimental design minimized as much as possible any effect of cold temperature endurance - in the sense of the ability to survive at 10°C - on our phenotype.

We apologize for not clarifying that point earlier and have added this text to the Results: “We chose 5 weeks based on pilot studies that showed that nearly all DGRP lines showed excellent survival at 5 weeks in diapause conditions while exhibiting sufficient variation in post-diapause fecundity to carry out GWAS. Beyond 5 weeks, fecundity was low and there was insufficient variation to conduct a GWAS.”

Additionally, I raised concerns about the validity of prioritizing genes with multiple associated variants. Although the authors agreed with this point, they did not revise the manuscript accordingly. The statement that 'Genes with multiple SNPs are good candidates for influencing diapause traits' is not a valid argument within the context of population and quantitative genetics.

We apologize for neglecting to revise the manuscript accordingly. We have revised Supplemental Table: S4 and ranked the genes by p-value.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The effects of environmental stress on animal life are gaining importance with climate change. Diapause is a dormancy program that occurs in response to an adverse environment, followed by resumption of development and reproduction upon the return of favorable conditions. Diapause is a complex trait, so we leveraged the Drosophila genetic reference panel (DGRP) lines and conducted a Genome-Wide Association Study (GWAS) to characterize the genetic basis of diapause. We assessed post-diapause and non-diapause fecundity across 193 DGRP lines. GWAS revealed 546 genetic variants, encompassing single nucleotide polymorphisms, insertions and deletions associated with post-diapause fecundity. We identified 291 candidate diapause-associated genes, 40 of which had previously been associated with diapause. 89 of the candidates were associated with more than one SNP. Gene network analysis indicated that the diapause-associated genes were primarily linked to neuronal and reproductive system development. Similarly, comparison with results from other fly GWAS revealed the greatest overlap with olfactory-behavior-associated and fecundity-and-lifespan-associated genes. An RNAi screen of the top candidates identified two neuronal genes, Dip-γ and Scribbler, to be required during recovery for post-diapause fecundity. We complemented the genetic analysis with a test of which neurons are required for successful diapause. We found that although amputation of the antenna had little to no effect on non-diapause lifespan, it reduced diapause lifespan and postdiapause fecundity. We further show that olfactory receptor neurons and temperature-sensing neurons are required for successful recovery from diapause. Our results provide insights into the molecular, cellular, and genetic basis of adult reproductive diapause in Drosophila.

Revised version: The effects of environmental stress on animal life are gaining importance with climate change. Diapause is a dormancy program that occurs in response to an adverse environment, followed by resumption of development and reproduction upon the return of favorable conditions. Diapause is a complex trait, so we leveraged the Drosophila genetic reference panel (DGRP) lines and conducted a Genome-Wide Association Study (GWAS) to characterize the genetic basis of diapause. We assessed post-diapause and non-diapause fecundity across 193 DGRP lines. GWAS revealed 546 genetic variants, encompassing single nucleotide polymorphisms, insertions and deletions associated with post-diapause fecundity. We identified 291 candidate diapause-associated genes, 40 of which had previously been associated with diapause, and 89 of which were associated with more than one SNP. Gene network analysis indicated that the diapause-associated genes were primarily linked to neuronal and reproductive system development. Similarly, comparison with results from other fly GWAS revealed the greatest overlap with olfactory-behavior-associated and fecundity-and-lifespan-associated genes. An RNAi screen of selected candidates identified two neuronal genes, Dip-𝛾 and Scribbler, to be required during recovery for post-diapause fecundity. We complemented the genetic analysis with a test of which neurons are required for successful diapause. We found that although amputation of the antenna had little to no effect on non-diapause lifespan, it reduced diapause lifespan and postdiapause fecundity. We further show that olfactory receptor neurons and temperature-sensing neurons are required for successful recovery from diapause. Our results provide insights into the molecular, cellular, and genetic basis of adult reproductive diapause in Drosophila.

Introduction

To endure adverse environments, many organisms from worms and flies to some mammals initiate a dormancy program called diapause, triggered by cues like unfavorable temperatures, drought, or starvation (Deng et al., 2018; Easwaran and Montell, 2023; Hand et al., 2016; Hussein et al., 2022; Tatar and Yin, 2001). By enhancing survival, diapause can support both beneficial insects like pollinators, and harmful insects such as agricultural pests and disease vectors. Elucidating diapause mechanisms may suggest ways to control harmful insects and support beneficial organisms. Furthermore, unraveling diapause genetics holds promise for understanding stem cell resilience (Easwaran et al., 2022), aging (Fontana et al., 2010; Hutfilz, 2022), and even cancer dormancy (Easwaran and Montell, 2023; Lin and Zhu, 2021).

Successful diapause occurs optimally at different developmental stages depending on the organism (Denlinger, 2002; Hahn and Denlinger, 2011; Kostál, 2006). In Drosophila species, the optimal stage is the newly-eclosed adult. During diapause, feeding, activity levels, and mating behaviors change, metabolism slows, reproduction arrests, stress resilience increases, and lifespan is lengthened (Hutfilz, 2022; Kubrak et al., 2016, 2014; Reis et al., 2015; Tatar and Yin, 2001; Tatar et al., 2001; Zonato et al., 2017).

The genetic basis of adult reproductive diapause in Drosophila is likely complex because the diapause program involves sensing environmental cues and responding by reprogramming behavior, metabolism, reproduction, and aging, continuously monitoring environmental conditions, and reactivating development and/or reproduction when conditions become favorable. Historically, studies of Drosophila diapause have focused primarily on entry into diapause by measuring arrest of egg production, specifically at the stage of yolk accumulation referred to as vitellogenesis (Saunders et al., 1990, 1989). However, it has become clear that diapause is a comprehensive program that changes virtually every aspect of the fly’s life from behavior and to metabolism, reproduction, and lifespan (Denlinger, 2023; Kubrak et al., 2014; Tatar and Yin, 2001). Prior studies have identified temperature and day-length as relevant environmental cues and key hormones such as insulin, juvenile hormone (JH), and 20- hydroxyecdysone (20E), which coordinate metabolism and reproduction (Denlinger, 2023; Flatt et al., 2008; Hutfilz, 2022; Saunders et al., 1990; Sim and Denlinger, 2013; Tatar and Yin, 2001; Tatar, 2004; Tatar et al., 2001; Williams et al., 2006). Recent studies identify specific circadian neurons that sense changes in temperature and relay that information to the reproductive system (Hidalgo et al., 2023; Meiselman et al., 2022). Yet, much remains to be learned about this fascinating life history trait (Denlinger, 2023).

One approach to analyzing complex traits and mapping quantitative trait loci is to conduct a Genome-Wide Association Study (GWAS) (Huang et al., 2014; Morozova et al., 2015; Shorter et al., 2015). The Drosophila Genome Reference Panel (DGRP) comprises fully sequenced, highly inbred lines of Drosophila melanogaster and serves as a valuable tool for understanding genotype-phenotype relationships in Drosophila (Mackay and Huang, 2018; Mackay et al., 2012; Morozova et al., 2015).

We used the DGRP to carry out a Drosophila GWAS and identify genes and gene classes associated with successful diapause. We developed a novel method for assessing successful diapause, by evaluating not only the arrest of egg development or the recovery of egg maturation but a more stringent criterion: the ability of post-diapause flies to produce viable adult progeny. Of the 291 diapause-associated genes we identified, ∼40 had previously been implicated in diapause in flies or another organism. The classes of genes most over-represented in the diapause-associated set are genes that regulate neuronal development and gonad development. We show that Gal4-mediated RNAi knockdown is effective at 10°C, and we identify two neuronal genes required during recovery for post-diapause fecundity. We further complement this genetic analysis by identifying neurons in the fly antenna required for successful diapause. Our studies implicate the olfactory system as critically important in successful diapause.

Results

Quantifying successful diapause recovery across the DGRP lines

Earlier version: Key features of Drosophila diapause include ovarian arrest and post-diapause recovery of fertility (the ability to produce any progeny) and fecundity (the number of progeny produced). Whereas the majority of studies of Drosophila diapause focus on arrest of oogenesis, we quantified diapause recovery by assessing the ability of newly-eclosed flies to undergo 35 days of diapause, recover, and produce viable progeny, which is a more stringent test of success (Fig. 1A).

Revised version: Key features of Drosophila diapause include ovarian arrest and post-diapause recovery of fertility (the ability to produce any progeny) and fecundity (the number of progeny produced). Whereas the majority of studies of Drosophila diapause focus on arrest of oogenesis, we chose to quantify diapause recovery by assessing the ability of newly-eclosed flies to undergo 35 days of diapause, recover, and produce viable progeny, which is a more stringent test of success (Fig. 1A). Additionally, in pilot experiments, post-diapause fecundity exhibited sufficient variation to be useful for a GWAS.

Quantification of diapause in Drosophila Genetic Reference Panel (DGRP) lines measured as the ratio of post-diapause to non-diapause fecundity.

(A) Schematic of experimental workflow. (B) Average number of progenies produced (fecundity) in a 4-day individual female fly mating experiment of DGRP lines either as non-diapausing (yellow) or after a 35-day post-diapause (blue) virgin flies. Each dot represents the average fecundity and the line above represents the standard error. DGRP line numbers are indicated wherever the post- diapause fecundity exceeds the non-diapause fecundity. (C) Normalized post-diapause fecundity [average of (individual post-diapause fecundity/mean non-diapause fecundity)] of each DGRP line. (D) Correlation of post-diapause to non-diapause fecundity. Pearson’s r=0.5682.r2=0.3228. (E-F) Frequency distribution of DGRP lines fecundity under non-diapause (E) and of the normalized post-diapause fecundity (F). (G) Average 4-day fecundity of single female flies, each crossed with 2 young Canton S male flies, aged for 1, 35, or 42 days in non- diapause conditions or kept in diapause conditions for 35 or 42 days followed by recovery. 1- way ANOVA and Tukey’s multiple comparison test, compact letter display shows comparisons. n is the number of individual female fly fecundity measured.

Added in the revised version: DGRP lines (n=193) were allowed to develop under non-diapause conditions (25°C and 12:12 L:D). Virgin females were collected and transferred to diapausing conditions (10°C and 8:16 L:D) for 5 weeks. Subsequently, flies were shifted to 18°C for one day, followed by incubation at 25°C in a fresh vial of fly food for recovery. We chose 5 weeks based on pilot studies that showed that nearly all DGRP lines showed excellent survival at 5 weeks in diapause conditions while exhibiting sufficient variation in post-diapause fecundity to carry out GWAS. Beyond 5 weeks, fecundity was low and there was insufficient variation for a GWAS.

Earlier version: DGRP lines were allowed to develop under non-diapause conditions (25°C and 12:12 L:D). Virgin females were collected and transferred to diapausing conditions (10°C and 8:16 L:D) for 5 weeks. Subsequently, flies were shifted to 18°C for one day, followed by incubation at 25°C in a fresh vial of fly food for recovery. We evaluated the ability of post-diapause flies to reproduce by mating individual females with 2 Canton-S (CS) male control (non-diapausing) flies for 4 days. Afterward, parents were removed, and progeny were allowed to develop for 12 more days. We then measured post-diapause fecundity by counting the number of adult progeny that eclosed. A minimum of 20 female flies per DGRP line were tested. Non-diapause fecundity was measured by setting crosses immediately upon collecting virgin flies. The average non-diapause and post-diapause 4-day fecundity of DGRP lines are shown in Figure 1B, arranged in ascending order of non-diapause fecundity. In general, most lines showed lower fecundity post-diapause compared to newly-eclosed flies in optimal conditions, as previously reported for one inbred wild strain (Tatar et al., 2001). However, there were interesting exceptions. For example, lines 309, 738, and 853 exhibited better fecundity post-diapause compared even to young flies in optimal conditions (Supplemental Table S1). The DGRP lines exhibit variability in fecundity both post-diapause and in non-diapause conditions [Fig. 1B and (Durham et al., 2014)]. Therefore, we normalized post-diapause to non-diapause fecundity (Fig. 1C and Supplemental Table S1).

Revised version: We evaluated the ability of post-diapause flies to reproduce by mating individual females with 2 Canton-S (CS) male control (non-diapausing) flies for 4 days. Afterward, parents were removed, and progeny were allowed to develop for 12 more days. We then measured post-diapause fecundity by counting the number of adult progeny that eclosed. A minimum of 20 female flies per DGRP line were tested. Non-diapause fecundity was measured by setting crosses immediately upon collecting virgin flies. The average non-diapause and post-diapause 4-day fecundity of DGRP lines are shown in Figure 1B, arranged in ascending order of non- diapause fecundity. In general, most lines showed lower fecundity post-diapause compared to newly-eclosed flies in optimal conditions, as previously reported for one inbred wild strain (Tatar et al., 2001). However, there were interesting exceptions. For example, lines 309, 738, and 853 exhibited better fecundity post-diapause compared even to young flies in optimal conditions (Supplemental Table S1). The DGRP lines exhibit variability in fecundity both post- diapause and in non-diapause conditions [Fig. 1B and (Durham et al., 2014)]. Therefore, we normalized post-diapause to non-diapause fecundity (Fig. 1C and Supplemental Table S1). The broad-sense heritability for normalized post-diapause fecundity was 0.51 (see Methods).

The lines that showed poor non-diapause fecundity that improved after diapause (e.g.362, 737, 822) suggested that there might be a genetic tradeoff between the two. However, the overall correlation (Fig. 1D) between non-diapause and post-diapause fecundity was 32.3% (R2=0.3228, Pearson’s correlation coefficient, r=0.5682), suggesting a positive rather than negative correlation overall between post-diapause and non-diapause fecundity. The frequency distributions of 4-day fecundity for the non-diapause (Fig. 1E), and normalized fecundity (Fig.1F) conditions are shown. Although fecundity is lower post-diapause than in young flies, even in Canton-S (CS) controls, it is higher than for flies maintained in non-diapause conditions for 35 or 42 days (Fig. 1G and Supplemental Table S2).

GWAS for diapause using the DGRP tool

The normalized fecundity scores (average of normalized post-diapause fecundity (individual post-diapause 4-day fecundity / average non-diapause 4-day fecundity)] served as the basis for conducting a GWAS using the DGRP2 web tool (Huang et al., 2014; Mackay et al., 2012). A total of 546 genetic variants, encompassing single nucleotide polymorphisms (SNPs) and insertions/deletions (indels), were identified as associated with diapause fecundity (Fig. 2A-B and Supplemental Table S3). When a variant is located within 1kb up- or down-stream of an annotated gene, or within the gene, it is considered potentially associated with that gene. We thus identified 291 candidate diapause-associated fecundity genes. Notably, 40 out of the 291 genes had previously been reported to be associated with diapause (Supplemental Table S4) either through functional analysis [e.g., insulin receptor (InR) and couch potato (cpo)] (Kankare et al., 2012; Kubrak et al., 2014; Schmidt et al., 2008; Sim and Denlinger, 2013; Zhang and Denlinger, 2011), due to changes in gene expression [e.g., expanded (ex) and Laminin A (LanA)] (Zhao et al., 2016), or associated with cold tolerance [β-Tub97EF (Myachina et al., 2017)] (Supplemental Table S4). Eighty-nine genes were associated with more than one diapause- associated variant (e.g. at least two different SNPs) (Fig. 2C).

Genome-wide association of Drosophila diapause.

(A) Manhattan plot for genome- wide association distribution. The position of each point along the y-axis indicates -log10(P- value) of association of a SNP, insertion, or deletion. Points above the blue line have a P value < 1e-5. The red line represents Bonferroni corrected P value = 4.8e-8. (B) Q-Q plot of P values from the DGRP single variant GWAS with the red line representing expected P-value and observed P values deviating (black dots ) from expected. (C) Numbers of genetic variants and candidate genes associated with diapause according to the GWAS. (D) Subnetworks from the Cytoscape analysis showing q-value (using the Benjamini-Hochberg procedure) for each subnetwork identified.

Network analysis of diapause-associated genes

Using the gene list obtained from the GWAS, we performed a gene network analysis using the GeneMania application within Cytoscape (Shannon et al., 2003; Warde-Farley et al., 2010). This application facilitates the generation of network predictions based on a combination of known physical and genetic interactions, co-localization, co-expression, shared protein domains, pathway data, and predicted functional relationships between genes. Additionally, it identifies subnetworks of genes related by functional Gene Ontology (GO). Enriched subnetworks are identified by dividing the number of genes from the input set by the total number of genes associated with a specific gene ontology. The subnetwork analysis is presented in Figure 2D. Two primary classes of GO terms are significantly enriched in the diapause GWAS set compared to the genome as a whole: nervous system development and the reproductive system/gonad development. These two categories are consistent with the idea that changes in the environment are sensed by the nervous system and the information is relayed to the reproductive system.

We also compared the list of diapause-associated genes to results of other Drosophila GWAS (Fig. 3, Supplemental Table S5). Intriguingly, the most significant overlap was with genes identified as associated with olfactory behavior (Horváth and Kalinka, 2018), which has not previously been described as important for diapause. The second most significant overlap was with genes identified in non-diapause fecundity and lifespan (Durham et al., 2014). Less significant overlaps were found with genes associated with circadian rhythm (Harbison et al., 2019), chill coma recovery (Mackay et al., 2012), development time associated with lead toxicity (Zhou et al., 2016), alcohol sensitivity (Morozova et al., 2015), starvation resistance (Mackay et al., 2012), mean food intake (Garlapow et al., 2015), sleep (Harbison et al., 2013), neurotoxin methylmercury (MeHg) tolerance (Montgomery et al., 2014), brain wiring regulators (Li et al., 2020), and death due to traumatic brain injury (Katzenberger et al., 2015). Overlaps with genes associated with courtship, viability in lead toxicity (Zhou et al., 2016), effects of genetic architecture and Wolbachia status on starvation (Huang et al., 2014), aggression (Shorter et al., 2015), ER stress (Chow et al., 2013), leg development (Grubbs et al., 2013), startle response (Mackay et al., 2012), and embryo development time (Horváth et al., 2016) were not statistically significant. In summary, both analyses implicated the nervous system, and potentially the olfactory system in particular, in the regulation of post-diapause fecundity.

Common genes to diapause-GWAS hits and other behavior-associated genes.

(A-T) Venn diagrams illustrate the intersection of genes associated with diapause identified through Genome-Wide Association Studies (diapause-GWAS) with genes from other behavior-related gene lists obtained from various studies. The diapause-GWAS gene set is represented as the first set throughout the figure, while subsequent sets represent different behavior-related gene lists identified in separate studies. The percentage of common genes compared to the total genes from different respective behavior-associated gene lists are provided for each Venn diagram. p-values of overlap to the diapause gene list determined by Fisher’s exact tests are also provided. Venn diagrams are arranged in the order of p-values.

RNAi analysis of GWAS candidates identifies neural genes required for recovery

Earlier version: Genes with multiple SNPs are good candidates for influencing diapause traits. To assess the functional significance of the top candidates, we conducted an RNAi screen of the genes with multiple associated SNPs. First, we evaluated the effectiveness of Gal4 under diapausing conditions. As a control, we crossed Mat-α-tub-Gal4, which drives expression in the female germline to UASp-F-tractin.tdTomato to assess the effectiveness of Gal4-mediated expression at 10°C. We compared tdTomato expression in flies maintained at different temperatures for 3 weeks. Compared to 25°C (Fig. 4A) and 18°C (Fig. 4B), flies at 10°C exhibited as high or higher expression of tdTomato (Fig. 4C).

Revised version: To assess the functional significance of the top candidates, we conducted an RNAi screen. First, we evaluated the effectiveness of Gal4 under diapausing conditions. As a control, we crossed Mat-α-tub-Gal4, which drives expression in the female germline to UASp-F-tractin.tdTomato to assess the effectiveness of Gal4-mediated expression at 10°C. We compared tdTomato expression in flies maintained at different temperatures for 3 weeks. Compared to 25°C (Fig. 4A) and 18°C (Fig. 4B), flies at 10°C exhibited as high or higher expression of tdTomato (Fig. 4C). To evaluate the effectiveness of Gal4-mediated RNAi at 10°C, we used Mat-α-tub-Gal4 to drive zpg (zero population growth, aka INX4) RNAi and assessed the extent of Zpg knockdown by immunostaining with an anti-Zpg antibody. We used Mat-α-tub-Gal4 because Zpg is expressed in the germline, and we chose to target Zpg because of the availability of the anti-Zpg antibody. Zpg is expressed from the earliest stages of germline development in the germarium (Fig. 4E, E’). Mat-α-tub-Gal4 is not expressed detectably in the germarium but is expressed in early egg chambers (Fig. 4A-C). So, to quantify the knockdown efficiency, we measured the level of Zpg in stage 3 egg chambers relative to the germarium staining, as an internal control. Remarkably, comparison of Mat-α-tub-Gal4>zpg RNAi flies kept at different temperatures for 3 weeks revealed that Gal4-mediated RNAi was as effective at 10°C, as it was at 18°C and 25°C (Fig. 4D-H’).

RNAi-mediated loss of function study to identify genes involved in diapause.

(A-C) Mat-ɑ-tub-Gal4 driving expression of UASp-F-tractin.tdTomato (red) at the indicated temperatures. Scale bars are 100µm. (D) Quantification of zpg RNAi knockdown in Stage-3 egg chambers normalized to the level of Zpg in the germarium. (1-way ANOVA and Tukey’s multiple comparison test, compact letter display shows comparisons). The numbers (n) of Stage 3 egg chambers quantified are shown. (E-H) Representative images of egg chambers stained with anti-Zpg antibody (green) from either control (no-knockdown), (E, E’) or knockdown of zpg (F- H’) driven by Mat-ɑ-tub-Gal4 at different temperatures. (E’-H’) are higher magnification, single channel views of the ovarioles shown in (E-H). Scale bars are 100 µm for (E-H) and 20 µm in (E’- H’). All flies in (A-H’) were kept at respective temperatures for 3 weeks. (I) Experimental design for RNAi knockdown specifically during recovery for the experiment shown in (J). The temperature-sensitive Gal80 repressor of Gal4 prevented RNAi expression during development and diapause. Incubation at 30°C during recovery inactivates Gal80, allowing Gal4-mediated RNAi knockdown. (J) Ubiquitous knockdown of Dip-𝛾 or sbb with tub Gal4 specifically during recovery as shown in (I) significantly reduces post-diapause/non-diapause fecundity compared to the control (tubGal80TS;tubGal4 > Ctrl RNAi #9331). (K) Pan-neuronal RNAi knockdown of Dip-𝛾 and sbb with nSybGal4 significantly reduces post-diapause/non-diapause fecundity compared to the control (nSyb Gal4> Ctrl RNAi #54037). (L) Glia-specific knockdown of Dip-𝛾 or sbb with Repo Gal4 causes little or no reduction in post-diapause/non-diapause fecundity (Control- Repo Gal4> Ctrl RNAi #54037). In (J-L), 1-way ANOVA and Tukey’s multiple comparison test, compact letter display shows comparisons. n is the number of individual female flies tested.

Removed in the revised version: To evaluate the effectiveness of Gal4-mediated RNAi at 10°C, we used Mat-α-tub-Gal4 to drive zpg (zero population growth, aka INX4) RNAi and assessed the extent of Zpg knockdown by immunostaining with an anti-Zpg antibody. We used Mat-α-tub-Gal4 because Zpg is expressed in the germline, and we chose to target Zpg because of the availability of the anti-Zpg antibody. Zpg is expressed from the earliest stages of germline development in the germarium (Fig. 4E, E’). Mat-α-tub-Gal4 is not expressed detectably in the germarium but is expressed in early egg chambers (Fig. 4A-C). So, to quantify the knockdown efficiency, we measured the level of Zpg in stage 3 egg chambers relative to the germarium staining, as an internal control. Remarkably, comparison of Mat-α-tub-Gal4>zpg RNAi flies kept at different temperatures for 3 weeks revealed that Gal4-mediated RNAi was as effective at 10°C, as it was at 18°C and 25°C (Fig. 4D-H’).

Earlier version: The top 15 GWAS hits for which multiple RNAi lines were available were selected for RNAi screening (Supplemental Table S6). We used tub-Gal4 because it is widely expressed, including in neurons. We measured non-diapause and post-diapause fecundity. For the five RNAi lines that proved lethal, we used tubulin-Gal80ts to suppress the RNAi expression during development and diapause and then shifted the flies to 30 °C to inactivate the Gal80 and drive the RNAi during recovery (Fig. 4I), and measured post-diapause fecundity. In this way we could test specifically for a function in post-diapause recovery (Supplemental Table S5). RNAi against two genes, Defective-proboscis-extension-response interacting protein-γ (Dip-γ) and Scribbler (sbb) significantly reduced post-diapause fecundity (Fig. 4J).

Revised version: Fifteen GWAS hits for which multiple RNAi lines were available were selected for RNAi screening (Supplemental Table S6). We used tub-Gal4 because it is widely expressed, including in neurons. We measured non-diapause and post-diapause fecundity. For the five RNAi lines that proved lethal, we used tubulin-Gal80ts to suppress the RNAi expression during development and diapause and then shifted the flies to 30 °C to inactivate the Gal80 and drive the RNAi during recovery (Fig. 4I), and measured post-diapause fecundity. In this way we could test specifically for a function in post-diapause recovery (Supplemental Table S5). RNAi against two genes, Defective-proboscis-extension-response interacting protein-𝛾 (Dip-𝛾) and Scribbler (sbb) significantly reduced post-diapause fecundity (Fig. 4J).

To determine whether Dip-𝛾 and sbb are required specifically in neurons we crossed nSyb-Gal4, a pan-neuronal driver, to validated UAS-Dip-𝛾 RNAi and UAS-sbb RNAi lines (Davis et al., 2014; Shimozono et al., 2019). Neuron-specific knockdown of Dip-𝛾 (Fig. 4K, Supplemental Table S7) caused as severe a defect in post-diapause fecundity as ubiquitous RNAi, in contrast to glial knockdown using Repo-Gal4 (Fig. 4L, Supplemental Table S7). Pan neuronal RNAi of sbb (Fig. 4K) also significantly inhibited post-diapause fecundity more than glial RNAi (Fig. 4L). We conclude that Dip-𝛾 and sbb are required in neurons for successful post-diapause fecundity, consistent with the enrichment of this gene class in the diapause GWAS.

Post-diapause fecundity requires neurons in the antenna

While we were able to show significant functional effects of Dip-𝛾 and sbb, GWAS studies by their nature identify many genes with small effects, which are too small to detect individually. Furthermore, the overlap with genes associated with olfactory behavior led us to complement the genetic analysis by testing whether the antenna and neurons within it are required for successful diapause. We removed the antenna from CS flies and measured post-diapause fecundity. As a control for the surgery, we removed the arista, which is an appendage from the antenna. Removal of the antenna but not the arista reduced post-diapause fecundity compared to unmanipulated controls (Fig. 5A and Supplemental Table S8). Removal of the antenna but not the arista also reduced the number of germline stem cells post-diapause (Fig. 5B). These results suggest that sensory cells in the antenna are important for successful recovery post- diapause.

Neural control of diapause.

(A) Effect of antenna removal on recovery of fecundity post-diapause. Arista removal was used as a control for the surgery. n is the number of individual female flies tested. 1-way ANOVA and Tukey’s multiple comparison test, compact letter display shows comparisons. (B) Effect of antenna removal on GSC recovery after 5 weeks of diapause. n is the number of germaria counted (there are typically 2-3 GSCs/germarium). 1- way ANOVA and Tukey’s multiple comparison test with compact letter display to show comparisons. (C) Role of antenna in lifespan extension in diapause. Control flies were maintained at 25°C and diapause flies were moved to 10°C. Median survival for flies with intact antenna in diapause (+A_Diap) - 142.5 days; antennaless flies in diapause (-A_Diap) - 7 days; Control with intact antenna in optimal conditions (+A_Ctrl) - 75 days; and antennless flies in optimal conditions (-A_Ctrl) - 72 days. Survival curves are compared pairwise using the Log-rank (Mantel-Cox) test and ꭓ2 values are: Diap +/- antenna = 98.89, Ctrl +/- antenna = 11.69, Diap +antenna vs Ctrl +antenna = 163.5, and Diap -antenna vs Ctrl -antenna = 0.5. (D) Control (+A_Diap) and antennaless (-A_Diap) flies were maintained at 25°C for 2 weeks post-surgery to allow for wound healing and shifted to the diapausing conditions. Median survival for +A_Diap- 142 days; -A_Diap - 95 days. Survival curves are compared pairwise using the Log-rank (Mantel-Cox) test and ꭓ2 values =86. (E-F) Effects on diapause of inactivating neuronal transmission (by driving tetanus toxin using UAS-TeTxLC.tdt, (E) in odor responsive neurons using the indicated Gal4 lines. Orco is the co-receptor for all of the odorant receptors. (F) Ir8a is a co-receptor involved in organic acid detection. Ir25a is a co-receptor involved in chemo- and thermo-sensation. Ir76b is involved in detection of various amines and salt. Ir84a is involved in detection of phenylacetic acid and male courtship behavior. Hot Cells are heat-sensitive cells in the arista.

Another important diapause trait is lifespan extension. So, we tested the effect of removing the antenna on post-diapause and non-diapause lifespan (Fig. 5C and D). Antenna removal had little effect on non-diapause lifespan (median survival 72 days w/o antenna vs 75 days with antenna) (Fig. 5C). In contrast, antennaless flies exhibited a substantial reduction in diapause lifespan compared to flies with antennae (Fig. 5C). We considered the possibility that the precipitous deaths of many antennaless flies at 10 °C (Fig. 5C) might be a consequence of impaired wound healing. So, we repeated the experiments, allowing the flies 2 weeks to recover at 25 °C prior to moving them to the diapause-inducing temperature of 10 °C (Fig. 5D). Two weeks at 25 °C substantially prevented the early post-surgery death. Yet, antennaless flies still had shorter lifespans (median survival 95 days) compared to intact flies (median survival 142 days), supporting the importance of the antenna in lifespan extension (see discussion). A potential caveat for the latter experiment is that two weeks at 25 °C might impair diapause induction. However, flies with antennae exhibited very similar lifespan extension at 10 °C with or without the 2 week wound-healing period at 25 °C (compare dark green lines in Fig. 5C and D), mitigating that concern. We conclude that the antennae are likely important for lifespan extension and post-diapause fecundity.

Earlier version: Drosophila antennae are involved in various sensory modalities, including olfaction. To identify which neurons are required for diapause, we suppressed neuronal activity in subsets of cells by expressing the tetanus toxin light chain protein (TNT), which selectively cleaves the neuronal isoform of fly synaptobrevin. Orco is the virtually universal co-receptor for all of the odorant receptors of the OR class (Montell, 2021), so we used Orco-Gal4 to drive expression. We found that blocking neuronal transmission in the Orco-expressing neurons decreased the post-diapause/non-diapause fecundity ratio (Fig. 5E, Supplemental Table S9), further implicating odor perception in successful diapause recovery (see discussion).

Revised version: Drosophila antennae are involved in various sensory modalities, including olfaction. To identify which neurons are required for diapause, we suppressed neuronal activity in subsets of cells by expressing the tetanus toxin light chain protein (TNT), which selectively cleaves the neuronal isoform of fly synaptobrevin. Orco is the virtually universal co-receptor for all of the odorant receptors of the OR class (Montell, 2021), so we used Orco-Gal4 to drive expression.

Added in the revised version: We found that blocking neuronal transmission in the Orco-expressing neurons decreased the post-diapause/non-diapause fecundity ratio (Fig. 5E, Supplemental Table S9), further implicating odor perception in successful diapause recovery (see discussion).

Ionotropic receptors (IRs) are another class of odorant receptor (Montell, 2021). Ir25a is broadly expressed, and inhibiting Ir25a-expressing neurons by expressing TNT reduced the post-diapause/non-diapause fecundity ratio compared to Gal4s expressed in other IR-expressing cells such as Ir8a and Ir84a, which are involved in odor perception, and Ir76b, which is involved in contact chemosensation (Fig. 5F and Supplemental Table S9). We also found that inhibiting neurons expressing Ir21a, mutations in which disrupt warm and cool avoidance behaviors (Budelli et al., 2019), dramatically impaired recovery from diapause. Blocking Ir21a neuronal transmission prevented flies from exiting diapause, as evidenced by 100% mortality of the post-diapause female flies moved to recovery conditions (n= 25). Expressing TNT in the “hot cell” neurons (Ni et al., 2013), which respond to heating, using HC-Gal4, also significantly reduced the post-diapause fecundity (Fig. 5F and Supplemental Table S9).

Discussion

Genome wide analysis of the effect of diapause on fecundity

Drosophila diapause is a fascinating life history trait that confers resilience in stressful environments and extends reproductive potential and organismal longevity. Our GWAS study identified nearly 300 candidate genes associated with an important but understudied feature of diapause: the ability to recover and reproduce successfully post-diapause. The most striking finding is that genes associated with neural development are highly over-represented in the diapause set. We further confirmed that at least two such genes, Dip-𝛾 and sbb, are essential for post-diapause fecundity.

When we started this project, little was known about the neural control of diapause. Consistent with our finding reported here that neural genes are enriched in the diapause GWAS, and our recent finding that circadian activity and sleep are dramatically altered in diapause (Meyerhof et al., 2024), two recent studies have reported that low temperature affects two subsets of circadian neurons, DN3s, and sLNvs, which in turn impact reproductive arrest (Hidalgo et al., 2023; Meiselman et al., 2022). sLNv neurons secrete the neuropeptide Pigment Dispersing Factor (PDF) onto insulin producing cells (IPCs), which secrete insulin-like peptides (ILPs), which in turn promote JH production. At 10 °C, PDF mRNA and protein levels are reduced in sLNv neurons, reducing JH production, which arrests vitellogenesis. In parallel, cool-temperature-induced reduction in DN3 activity disinhibits cholinergic ASTC-R2 neurons, which in turn inhibit vitellogenic egg chamber development. These neurons do not synapse onto the IPCs, and the mechanism by which the ASTC-R2 neurons affect oogenesis is unknown. Thus, there are at least two parallel mechanisms by which cool temperatures affect neural activity and contribute to reproductive arrest. It is not yet clear how cool temperatures reduce PDF mRNA or how PDF mRNA and protein levels rebound during recovery.

We found that the neuronal transcription factor Sbb is required post-diapause for recovery, so it is interesting to speculate that Sbb could affect PDF mRNA abundance. Sbb functions predominantly as a transcriptional repressor, so it could promote PDF expression in recovery by inhibiting a repressor of PDF. Alternatively, PDF reduces levels of the transcriptional co-activator Eyes Absent (Eya) in the IPCs, so reduced PDF leads to elevated Eya protein levels during diapause. It is possible that Sbb represses Eya transcription to facilitate recovery post- diapause. A third possibility is that Sbb functions in the DN3-ASTC-R2 pathway. We have also found that circadian rhythms and sleep behavior are dramatically altered in flies at 10°C (Meyerhof et al., 2024). It will be of interest to determine if Dip-𝛾 and sbb are required for those behavioral effects at low temperatures, in addition to the effects on post-diapause fecundity described here.

Olfactory neurons are required for post-diapause fecundity

We complemented our genetic analysis with an investigation of the cells that contribute to diapause. The overlap between olfactory-behavior-associated genes and diapause-associated genes inspired us to test if the antenna is required, and we found that the antenna is required both for post-diapause fecundity and for lifespan extension. We conclude that cells in the antenna transmit important information for successful diapause. Furthermore, we found that inhibiting olfactory receptor neurons by expressing TNT with Orco-Gal4 caused a similar effect on post-diapause fecundity as antenna removal. This is interesting because Meiselman et al found that the antenna was not required for accumulation of mature eggs post-diapause (Meiselman et al., 2022). Our measurement of fecundity required not only that mature eggs developed in the ovary but also that they were fertilized, laid, and could develop into adults. So, what is required for production of viable progeny in addition to mature eggs? Since both feeding and mating are suppressed during diapause, the results suggest that the antenna may be required to promote these behaviors, consistent with the known role of the antenna in providing sensory information important for flies to seek food and mates (Montell, 2021; Tunstall et al., 2012). The antenna may be dispensable for relieving the vitellogenesis block in the ovary but may specifically be required for animals to find sufficient food and mates post- diapause to support maximal fecundity.

Consistent with the possible role of the olfactory system in diapause traits, in a screen of 505 gene trap lines, Tunstall et al. identified 16 genes that were expressed in ORNs, three of which are present in our diapause-associated gene set: tai, cpo, and dpp (Tunstall et al., 2012). Cpo is a known diapause-associated gene (Kankare et al., 2012), and we show here that dpp is also required for post-diapause fecundity. Tunstall et al showed that tai is required for normal olfactory-mediated attraction to food. The olfactory system also regulates longevity in C. elegans (Alcedo and Kenyon, 2004) as well as Drosophila (Libert et al., 2007), and the precise effects depend on environmental factors including temperature and availability of food and mates (Allen et al., 2015), consistent with our findings on the role of olfaction in post-diapause lifespan extension and fecundity.

We further explored whether IR-expressing ORNs might also be required post-diapause by inhibiting subsets of IR-expressing cells with TNT. The most remarkable finding was that expressing TNT with Ir21a-Gal4 resulted in post-diapause death, indicating that Ir21a-Gal4- expressing cells are essential for successful exit from dormancy. Ir21a is an ionotropic receptor expressed in cells sensitive to cooling (cool cells). Animals mutant for Ir21a fail to detect cooling and also fail to avoid both warm and cool temperatures (Budelli et al., 2019). That cells required for temperature sensation are essential for diapause is perhaps not surprising, however it is at first glance puzzling that inhibiting the Ir21a-expressing cells affects diapause because Ir21a- expressing cells are located in the arista, which when amputated did not affect post-diapause fecundity in our experiments. Similarly puzzling is the observation that inhibition of the “hot cells,” which are also located in the arista in close proximity to the cool cells, also impairs post- diapause fecundity. One possible explanation for these results could be that cool cells and hot cells inhibit one another such that eliminating cool cell activity actually disinhibits hot cells and vice versa. In fact, Budelli et al (Budelli et al., 2019) found hot cell spiking in Ir21a mutants as well as defects in avoidance of both warm and cool temperatures. If inhibiting cooling-sensitive cells hyperactivates warming-sensitive cells and vice versa, then eliminating both cell types by removal of the arista might be less harmful than inactivating individual cell types.

In summary, we report the first GWAS of diapause, using the stringent assay of the ability to enter diapause, survive, exit diapause, and produce fertile offspring. This analysis has revealed a key role for the olfactory system, as well as specific genes and cells required for diapause lifespan extension and post-diapause fecundity. Additionally, the 291 candidate diapause-associated genes identified here provide many intriguing candidates for future study.

Methods

Drosophila stocks

Earlier version: The majority of the experiments were carried out using the inbred, sequenced lines of the Drosophila Genetic Reference Panel (DGRP) (Huang et al., 2014). Diapause trait quantified in 193 DGRP lines was used in the GWAS analysis. Please refer to The RNAi lines and Gal4 stocks were from BDSC and VDRC (GD44077, SH330386, 27049, GD10268, GD40698, 32034, KK101659, GD36166, 40889, 27508, GD991, GD992, GD21052, KK100412, 54811, GD51936, KK106911, 44661, 53315, 54037, 57840, KK104551, GD17767, 63013, 56867, 80461, KK104056, 35785, KK60102, 9331, 5138, 7019, 23292, 41731, 41728, 51311, 41750, 28838, 35607, 7063, 58989, 51635, 64349, and 7415). Hot cell-Gal4 (Gallio et al., 2011) is a gift from Zuker lab and Ir21a-Gal4 (Ni et al., 2016) is a gift from Craig Montell lab.

Revised version: The majority of the experiments were carried out using the inbred, sequenced lines of the Drosophila Genetic Reference Panel (DGRP) (Huang et al., 2014). Diapause trait quantified in 193 DGRP lines was used in the GWAS analysis. Please refer to Supplemental Table S1. The RNAi lines and Gal4 stocks were from BDSC and VDRC (GD44077, SH330386, 27049, GD10268, GD40698, 32034, KK101659, GD36166, 40889, 27508, GD991, GD992, GD21052, KK100412, 54811, GD51936, KK106911, 44661, 53315, 54037, 57840, KK104551, GD17767, 63013, 56867, 80461, KK104056, 35785, KK60102, 9331, 5138, 7019, 23292, 41731, 41728, 51311, 41750, 28838, 35607, 7063, 58989, 51635, 64349, and 7415). Hot cell-Gal4 (Gallio et al., 2011) is a gift from Zuker lab and Ir21a-Gal4 (Ni et al., 2016) is a gift from Craig Montell lab.

Fecundity assay

For the non-diapause fecundity assessment, a minimum of 15 newly-eclosed, not more than 6- hour-old virgin females of each strain were collected under mild CO2 anesthesia. These females were individually placed into fly food vials containing cornmeal, molasses, agar medium, and yeast. In individual female fly crosses, one female from the DGRP line and two male flies of the same genotype (Canton S) were added to each vial.

In the post-diapause group, a minimum of 20 newly eclosed, not more than 6-hour-old virgin female flies of each strain were collected and promptly transferred to a cold room under diapause conditions (10°C and 8L:16D) for 5 weeks. After this diapause induction period, the flies were shifted to 18°C and 12L:12D for 1 day for temperature acclimation. Subsequently, the flies were transferred to new vials with fresh food medium dusted with yeast and placed at 25°C and 12L:12D for an additional day. Similar to the non-diapause crosses, individual female fly crosses were set up. The flies were allowed to mate and lay eggs for 4 days at 25°C under a normal photoperiod of 12L:12D. After 4 days, the parental flies were discarded, and the resulting progeny were counted on the 16th day from the initial mating start date. The progeny count was limited to the number of adult flies and any black pupae produced by a single female during these 16 days. If female flies failed to survive diapause, the line was given the lowest score of zero, indicating an inability to diapause. Those who died during the 4-day cross period were not considered for scoring. In cases of low replicates due to deaths during diapause, recovery, or cross, the process was repeated to obtain sufficient replicates. The fecundity of flies can vary with the age of fly food, dryness of fly food, age of male flies used to assess the female fecundity, crowding in the bottle, etc. To control all these variables, we have always used fresh fly food (1-week-old) and 1-week-old male flies to reduce the variation. Also, we control the number of flies in collection bottles to 30-40 - about 25-30 females and 10 males and flipped every 3-4 days to a new bottle. We have also observed the fly vials regularly to prevent dryness of food and added water whenever necessary.

Similarly, assessments of non-diapause and post-diapause fecundity were conducted for the RNAi experiments. The post-diapause fecundity was normalized by dividing individual female fly fecundity by average non-diapause fecundity of the same genotype.

Scoring of fecundity

For both the non-diapause and the post-diapause, the fecundity of each strain was calculated based on the average number of progeny produced across all replicates, excluding cases in which the female and/or both males died during the 4-day mating window. To accurately assess the diapause capacity of each strain, we normalized it by dividing the individual post-diapause fecundity by the average of non-diapause fecundity for each DGRP line. The average of this normalized post-diapause fecundity were then utilized in our GWAS analysis pipeline (Mackay et al., 2012). This normalization eliminated the basal difference in fecundity among DGRP lines, allowing us to focus on the variability in diapause. We have also performed a one-way ANOVA to get the mean squares for between-group and within-group variances and calculated broad- sense heritability using the formula: H2 = MSG - MSE / MSG + (k-1) MSE where MSG - Mean square between groups and MSE - Mean square within groups and k - Number of individuals per group. Using this formula, the broad-sense heritability for normalized post-diapause fecundity was found to be 0.51.

Data analysis

GWAS was carried out using DGRP2 website. We conducted gene ontology enrichment and network analyses based on the top variants (P < 1e10−5) associated with the mean post- diapause fecundity/non-diapause fecundity score using the GeneMANIA application in Cytoscape (Shannon et al., 2003; Warde-Farley et al., 2010). The Gene Ontology (GO) categories and Q-values from the false discovery rate (FDR)-corrected hypergeometric test for enrichment are reported. Additionally, coverage ratios for the number of annotated genes in the displayed network versus the number of genes with that annotation in the genome are provided. GeneMANIA estimates Q-values using the Benjamini-Hochberg procedure.

Set Comparison Appyter was used for comparing the diapause-GWAS gene list to other behavior gene lists using 13500 as the background gene list and 0.05 as the significance level to calculate the p-value (Clarke et al., 2021).

Immunostaining

Added in the revised version: Immunofluorescence was performed using standard procedures. Briefly, adult female ovarioles were carefully dissected in phosphate-buffered saline (PBS) using a bent tungsten needle, pulling on the stalk region of older egg chambers to minimize damage to the germarium.

Earlier version: Immunofluorescence was performed using standard procedures. Briefly, adult female ovarioles were carefully dissected in phosphate-buffered saline (PBS) using a bent tungsten needle, pulling on the stalk region of older egg chambers to minimize damage to the germarium. Ovarioles prepared for immunostaining were fixed in 4% paraformaldehyde (PFA) in PBS for 20 minutes. To prevent sample sticking and facilitate settling, 20µl of PBST (PBS+0.2% Triton X100) was added to the fixing solution. After fixation, ovarioles were washed three times for 10 minutes each in PBST and then incubated with primary antibodies at 4°C overnight. The following morning, ovarioles were washed twice for 15 minutes in PBST before incubation with secondary antibodies for at least 2 hours at room temperature. After the removal of secondary antibodies, the samples were washed three times in PBST for 10 minutes each. Hoechst was used as a nuclear stain and added to the second wash solution. The ovarioles were subsequently mounted in Vectashield and stored at 4°C until imaging. All antibody dilutions were prepared in PBST. Due to difficulties in settling diapause samples in the solutions, a 5-minute hold step was introduced on a stand before removing solutions from the tube at each step, consistently followed for all controls as well. The anti-Zpg antibody (rabbit polyclonal) was used at a 1:20000 dilution (Smendziuk et al., 2015). The antibodies used for identifying Germline Stem Cells (GSCs) are Hts (1B1) and Vasa (DSHB) (Easwaran et al., 2022).

Revised version: Ovarioles prepared for immunostaining were fixed in 4% paraformaldehyde (PFA) in PBS for 20 minutes. To prevent sample sticking and facilitate settling, 20µl of PBST (PBS+0.2% Triton X100) was added to the fixing solution. After fixation, ovarioles were washed three times for 10 minutes each in PBST and then incubated with primary antibodies at 4°C overnight. The following morning, ovarioles were washed twice for 15 minutes in PBST before incubation with secondary antibodies for at least 2 hours at room temperature. After the removal of secondary antibodies, the samples were washed three times in PBST for 10 minutes each. Hoechst was used as a nuclear stain and added to the second wash solution. The ovarioles were subsequently mounted in Vectashield and stored at 4°C until imaging. All antibody dilutions were prepared in PBST. Due to difficulties in settling diapause samples in the solutions, a 5- minute hold step was introduced on a stand before removing solutions from the tube at each step, consistently followed for all controls as well. The anti-Zpg antibody (rabbit polyclonal) was used at a 1:20000 dilution (Smendziuk et al., 2015). The antibodies used for identifying Germline Stem Cells (GSCs) are Hts (1B1) and Vasa (DSHB) (Easwaran et al., 2022).

Lifespan analysis

Virgin female flies were collected and moved to the conditions to assess their lifespan. In the antenna removed fly lifespan measurement, before moving to the diapausing conditions we gave 14 days at optimal conditions (25°C, 12L:12D photoperiod) to heal the wound caused by amputation of the antenna. Flies were kept in food vials with a density of 20 or fewer flies per vial. Deaths were censored daily and all vials were flipped every other day for lifespan measurement at optimal condition (25°C, 12L:12D photoperiod) and once every month for lifespan measurement at diapause condition (10°C, 16L:8D photoperiod) to prevent desiccation and drying up of fly food. The Survival plot was created in Prism using the Kaplan-Meier survival analysis. Survival curves are compared pairwise by the Log-rank (Mantel-Cox) test to obtain the χ2 - values.